GPRC6A as a novel kokumi receptor responsible for enhanced taste preferences by ornithine

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

This paper contains what could be described as a "classic" approach towards evaluating a novel taste stimuli in an animal model, including standard behavioral tests (some with nerve transections), taste nerve physiology, and immunocytochemistry of taste cells of the tongue. The stimulus being tested is ornithine, from a class of stimuli called "kokumi" (in terms of human taste); these kokumi stimuli appear to enhance other canonical tastes, increasing what are essentially hedonic attributes of other stimuli. The mechanism for ornithine detection is thought to be GPRC6A receptors expressed in taste cells. The authors showed evidence for this in an earlier paper with mice; this paper evaluates ornithine taste in a rat model, and comes to a similar conclusion, albeit with some small differences between the two rodent species.

Strengths:

The data show effects of ornithine on taste/intake in laboratory rats: In two-bottle and briefer intake tests, adding ornithine results in higher intake of most, but all not all stimuli tested. Bilateral chorda tympani (CT) nerve cuts or the addition of GPRC6A antagonists decreased or eliminated these effects. Ornithine also evoked responses by itself in the CT nerve, but mainly at higher concentrations; at lower concentrations it potentiated the response to monosodium glutamate. Finally, immunocytochemistry of taste cell expression indicated that GPRC6A was expressed predominantly in the anterior tongue, and co-localized (to a small extent) with only IP3R3, indicative of expression in a subset of type II taste receptor cells.

Weaknesses:

As the authors are aware, it is difficult to assess a complex human taste with complex attributes, such as kokumi, in an animal model. In these experiments they attempt to uncover mechanistic insights about how ornithine potentiates other stimuli by using a variety of established experimental approaches in rats. They partially succeed by finding evidence that GPRC6A may mediate effects of ornithine when it is used at lower concentrations. In the revision they have scaled back their interpretations accordingly. A supplementary experiment measuring certain aspects of the effects of ornithine added to Miso soup in human subjects is included for the express purpose of establishing that the kokumi sensation of a complex solution is enhanced by ornithine; however, they do not use any such complex solutions in the rat studies. Moreover, the sample size of the human experiment is (still) small - it really doesn't belong in the same manuscript with the rat studies.

Despite the reviewer’s suggestion, we would like to include the human sensory experiment. Our rationale is that we must first demonstrate that the kokumi of miso soup is enhanced by the addition of ornithine, which is then followed by basic animal experiments to investigate the underlying mechanisms of kokumi in humans.

We did not present the additive effects of ornithine on miso soup in the present rat study because our previous companion paper (Fig. 1B in Mizuta et al., 2021, Ref. #26) already confirmed that miso soup supplemented with 3 mM L-ornithine (but not D-ornithine) was statistically significantly (P < 0.001) preferred to plain miso soup by mice.

Furthermore, we believe that our sample size (n = 22) is comparable to those employed in other studies. For example, the representative kokumi studies by Ohsu et al. (Ref. #9), Ueda et al. (Ref. #10), Shibata et al. (Ref. #20), Dunkel et al. (Ref. #37), and Yang et al. (Ref. #44) used sample sizes of 20, 19, 17, 9, and 15, respectively.

Reviewer #2 (Public review):

Summary:

The authors used rats to determine the receptor for a food-related perception (kokumi) that has been characterized in humans. They employ a combination of behavioral, electrophysiological, and immunohistochemical results to support their conclusion that ornithine-mediated kokumi effects are mediated by the GPRC6A receptor. They complemented the rat data with some human psychophysical data. I find the results intriguing, but believe that the authors overinterpret their data.

Strengths:

The authors provide compelling evidence that ornithine enhances the palatability of several chemical stimuli (i.e., IMP, MSG, MPG, Intralipos, sucrose, NaCl, quinine). Ornithine also increases CT nerve responses to MSG. Additionally, the authors provide evidence that the effects of ornithine are mediated by GPRC6A, a G-protein-coupled receptor family C group 6 subtype A, and that this receptor is expressed primarily in fungiform taste buds. Taken together, these results indicate that ornithine enhances the palatability of multiple taste stimuli in rats and that the enhancement is mediated, at least in part, within fungiform taste buds. This is an important finding that could stand on its own. The question of whether ornithine produces these effects by eliciting kokumi-like perceptions (see below) should be presented as speculation in the Discussion section.

Weaknesses:

I am still unconvinced that the measurements in rats reflect the "kokumi" taste percept described in humans. The authors conducted long-term preference tests, 10-min avidity tests and whole chorda tympani (CT) nerve recordings. None of these procedures specifically model features of "kokumi" perception in humans, which (according to the authors) include increasing "intensity of whole complex tastes (rich flavor with complex tastes), mouthfulness (spread of taste and flavor throughout the oral cavity), and persistence of taste (lingering flavor)." While it may be possible to develop behavioral assays in rats (or mice) that effectively model kokumi taste perception in humans, the authors have not made any effort to do so. As a result, I do not think that the rat data provide support for the main conclusion of the study--that "ornithine is a kokumi substance and GPRC6A is a novel kokumi receptor."

Kokumi can be assessed in humans, as demonstrated by the enhanced kokumi perception observed when miso soup is supplemented with ornithine (Fig. S1). Currently, we do not have a method to measure the same kokumi perception in animals. However, in the two-bottle preference test, our previous companion paper (Fig. 1B in Mizuta et al. 2021, Ref. #26) confirmed that miso soup supplemented with 3 mM L-ornithine (but not D-ornithine) was statistically significantly (P < 0.001) preferred over plain miso soup by mice.

Of the three attributes of kokumi perception in humans, the “intensity of whole complex tastes (rich flavor with complex tastes)” was partly demonstrated in the present rat study. In contrast, “mouthfulness (the spread of taste and flavor throughout the oral cavity)” could not be directly detected in animals and had to be inferred in the Discussion. “Persistence of taste (lingering flavor)” was evident at least in the chorda tympani responses; however, because the tongue was rinsed 30 seconds after the onset of stimulation, the duration of the response was not fully recorded.

It is well accepted in sensory physiology that the stronger the stimulus, the larger the tonic response—and consequently, the longer it takes for the response to return to baseline. For example, Kawasaki et al. (2016, Ref. #45) clearly showed that the duration of sensation increased proportionally with the concentration of MSG, lactic acid, and NaCl in human sensory tests. The essence of this explanation has been incorporated into the Discussion (p. 12).

Why are the authors hypothesizing that the primary impacts of ornithine are on the peripheral taste system? While the CT recordings provide support for peripheral taste enhancement, they do not rule out the possibility of additional central enhancement. Indeed, based on the definition of human kokumi described above, it is likely that the effects of kokumi stimuli in humans are mediated at least in part by the central flavor system.

We agree with the reviewer’s comment. Our CT recordings indicate that the effects of kokumi stimuli on taste enhancement occur primarily at the peripheral taste organs. The resulting sensory signals are then transmitted to the brain, where they are processed by the central gustatory and flavor systems, ultimately giving rise to kokumi attributes. This central involvement in kokumi perception is discussed on page 12. Although kokumi substances exert their effects at low concentrations—levels at which the substance itself (e.g., ornithine) does not become more favorable or (in the case of γ-Glu-Val-Gly) exhibits no distinct taste—we cannot rule out the possibility that even faint taste signals from these substances are transmitted to the brain and interact with other taste modalities.

The authors include (in the supplemental data section) a pilot study that examined the impact of ornithine on variety of subjective measures of flavor perception in humans. The presence of this pilot study within the larger rat study does not really mice sense. While I agree with the authors that there is value in conducting parallel tests in both humans and rodents, I think that this can only be done effectively when the measurements in both species are the same. For this reason, I recommend that the human data be published in a separate article.

Despite the reviewer’s suggestion, we intend to include the human sensory experiment. Our rationale is that we must first demonstrate that the kokumi of miso soup is enhanced by the addition of ornithine, and then follow up with basic animal experiments to investigate the potential underlying mechanisms of kokumi in humans.

In our previous companion paper (Fig. 1B in Mizuta et al., 2021, Ref. #26), we confirmed with statistical significance (P < 0.001) that mice preferred miso soup supplemented with 3 mM L-ornithine (but not D-ornithine) over plain miso soup. However, as explained in our response to Reviewer #2’s first concern (in the Public review), it is difficult to measure two of the three kokumi attributes—aside from the “intensity of whole complex tastes (rich flavor with complex tastes)”—in animal models.

The authors indicated on several occasions (e.g., see Abstract) that ornithine produced "synergistic" effects on the CT nerve response to chemical stimuli. "Synergy" is used to describe a situation where two stimuli produce an effect that is greater than the sum of the response to each stimulus alone (i.e., 2 + 2 = 5). As far as I can tell, the CT recordings in Fig. 3 do not reflect a synergism.

We appreciate your comments regarding the definition of synergy. In Fig. 5 (not Fig. 3), please note the difference in the scaling of the ordinate between Fig. 5D (ornithine responses) and Fig. 5E (MSG responses). When both responses are presented on the same scale, it becomes evident that the response to 1 mM ornithine is negligibly small compared to the MSG response, which clearly indicates that the response to the mixture of MSG and 1 mM ornithine exceeds the sum of the individual responses to MSG and 1 mM ornithine. Therefore, we have described the effect as “synergistic” rather than “additive.” The same observation applies to the mice experiments in our previous companion paper (Fig. 8 in Mizuta et al. 2021, Ref. #26), where synergistic effects are similarly demonstrated by graphical representation. We have also added the following sentence to the legend of Fig. 5:

“Note the different scaling of the ordinate in (D) and (E).”

Reviewer #3 (Public review):

Summary:

In this study the authors set out to investigate whether GPRC6A mediates kokumi taste initiated by the amino acid L-ornithine. They used Wistar rats, a standard laboratory strain, as the primary model and also performed an informative taste test in humans, in which miso soup was supplemented with various concentrations of L-ornithine. The findings are valuable and overall the evidence is solid. L-Ornithine should be considered to be a useful test substance in future studies of kokumi taste and the class C G protein coupled receptor known as GPRC6A (C6A) along with its homolog, the calcium-sensing receptor (CaSR) should be considered candidate mediators of kokumi taste. The researchers confirmed in rats their previous work on Ornithine and C6A in mice (Mizuta et al Nutrients 2021).

Strengths:

The overall experimental design is solid based on two bottle preference tests in rats. After determining the optimal concentration for L-Ornithine (1 mM) in the presence of MSG, it was added to various tastants including: inosine 5'-monophosphate; monosodium glutamate (MSG); mono-potassium glutamate (MPG); intralipos (a soybean oil emulsion); sucrose; sodium chloride (NaCl; salt); citric acid (sour) and quinine hydrochloride (bitter). Robust effects of ornithine were observed in the cases of IMP, MSG, MPG and sucrose; and little or no effects were observed in the cases of sodium chloride, citric acid; quinine HCl. The researchers then focused on the preference for Ornithine-containing MSG solutions. Inclusion of the C6A inhibitors Calindol (0.3 mM but not 0.06 mM) or the gallate derivative EGCG (0.1 mM but not 0.03 mM) eliminated the preference for solutions that contained Ornithine in addition to MSG. The researchers next performed transections of the chord tympani nerves (with sham operation controls) in anesthetized rats to identify a role of the chorda tympani branches of the facial nerves (cranial nerve VII) in the preference for Ornithine-containing MSG solutions. This finding implicates the anterior half-two thirds of the tongue in ornithine-induced kokumi taste. They then used electrical recordings from intact chorda tympani nerves in anesthetized rats to demonstrate that ornithine enhanced MSG-induced responses following the application of tastants to the anterior surface of the tongue. They went on to show that this enhanced response was insensitive to amiloride, selected to inhibit 'salt tastant' responses mediated by the epithelial Na+ channel, but eliminated by Calindol. Finally they performed immunohistochemistry on sections of rat tongue demonstrating C6A positive spindle-shaped cells in fungiform papillae that partially overlapped in its distribution with the IP3 type-3 receptor, used as a marker of Type-II cells, but not with (i) gustducin, the G protein partner of Tas1 receptors (T1Rs), used as a marker of a subset of type-II cells; or (ii) 5-HT (serotonin) and Synaptosome-associated protein 25 kDa (SNAP-25) used as markers of Type-III cells.

At least two other receptors in addition to C6A might mediate taste responses to ornithine: (i) the CaSR, which binds and responds to multiple L-amino acids (Conigrave et al, PNAS 2000), and which has been previously reported to mediate kokumi taste (Ohsu et al., JBC 2010) as well as responses to Ornithine (Shin et al., Cell Signaling 2020); and (ii) T1R1/T1R3 heterodimers which also respond to L-amino acids and exhibit enhanced responses to IMP (Nelson et al., Nature 2001). These alternatives are appropriately discussed and, taken together, the experimental results favor the authors' interpretation that C6A mediates the Ornithine responses. The authors provide preliminary data in Suppl. 3 for the possibility of co-expression of C6A with the CaSR.

Weaknesses:

The authors point out that animal models pose some difficulties of interpretation in studies of taste and raise the possibility in the Discussion that umami substances may enhance the taste response to ornithine (Line 271, Page 9).

Ornithine and umami substances interact to produce synergistic effects in both directions—ornithine enhances responses to umami substances, and vice versa. These effects may depend on the concentrations used, as described in the Discussion (pp. 9–10). Further studies are required to clarify the precise nature of this interaction.

One issue that is not addressed, and could be usefully addressed in the Discussion, relates to the potential effects of kokumi substances on the threshold concentrations of key tastants such as glutamate. Thus, an extension of taste distribution to additional areas of the mouth (previously referred to as 'mouthfulness') and persistence of taste/flavor responses (previously referred to as 'continuity') could arise from a reduction in the threshold concentrations of umami and other substances that evoke taste responses.

Thank you for this important suggestion. If ornithine reduces the threshold concentrations of tastants—including glutamate—and enhances their suprathreshold responses, then adding ornithine may activate additional taste cells. This effect could explain kokumi attributes such as an “extension of taste distribution” and possibly the “persistence of responses.” As shown in Fig. 2, the lowest concentrations used for each taste stimulus are near or below the thresholds, which indicates that threshold concentrations are reduced—especially for MSG and MPG. We have incorporated this possibility into the Discussion as follows (p.12):

“Kokumi substances may reduce the threshold concentrations as well as they increase the suprathreshold responses of tastants. Once the threshold concentrations are lowered, additional taste cells in the oral cavity become activated, and this information is transmitted to the brain. As a result, the brain perceives this input as coming from a wider area of the mouth.”

The status of one of the compounds used as an inhibitor of C6A, the gallate derivative EGCG, as a potential inhibitor of the CaSR or T1R1/T1R3 is unknown. It would have been helpful to show that a specific inhibitor of the CaSR failed to block the ornithine response.

Thank you for this important comment. We attempted to identify a specific inhibitor of CaSR. Although we considered using NPS-2143—a commonly used CaSR inhibitor—it is known to also inhibit GPRC6A. We agree that using a specific CaSR inhibitor would be beneficial and plan to pursue this in future studies.

It would have been helpful to include a positive control kokumi substance in the two bottle preference experiment (e.g., one of the known gamma glutamyl peptides such as gamma-glu-Val-Gly or glutathione), to compare the relative potencies of the control kokumi compound and Ornithine, and to compare the sensitivities of the two responses to C6A and CaSR inhibitors.

We agree with this comment. In retrospect, it may have been advantageous to directly compare the potencies of CaSR and GPRC6A agonists in enhancing taste preferences—and to evaluate the sensitivity of these preferences to CaSR and GPRC6A antagonists. However, we did not include γ-Glu-Val-Gly in the present study because we have already reported its supplementation effects on the ingestion of basic taste solutions in rats using the same methodology in a separate paper (Yamamoto and Mizuta, 2022, Ref. #25). The results from both studies are compared in the Discussion (p. 11).

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Major:

I am not convinced by the Author's arguments for including the human data. I appreciate their efforts in adding a few (5) subjects and improving the description, but it still feels like it is shoehorned into this paper, and would be better published as a different manuscript.

This human study is short, but it is complete rather than preliminary. The rationale for us to include the human data as supplementary information is shown in responses to the reviewer’s Public review.

Minor concerns:

Page 3 paragraph 1: Suggest "contributing to palatability".

Thank you for this suggestion. We have rewritten the text as follows:

“…, the brain further processes these sensations to evoke emotional responses, contributing to palatability or unpleasantness”.

Page 4 paragraph 2: The text still assumes that "kokumi" is a meaningful descriptor for what rodents experience. Re-wording the following sentence like this could help:

"Neuroscientific studies in mice and rats provide evidence that gluthione and y-Glu-Val-Gly activate CaSRs, and modify behavioral responses to other tastants in a way that may correspond to kokumi taste as experienced by humans. However, to our..."

Or something similar.

Thank you for this suggestion. We have rewritten the sentence according to your suggestion as follows:

"Neuroscientific studies (23,25,30) in mice and rats provide evidence that glutathione and y-Glu-Val-Gly activate CaSRs, and modify behavioral responses to other tastants in a way that may correspond to kokumi as experienced by humans”.

Page 7 paragraph 1 - put the concentrations of Calindol and EGCG used (in the physiology exps) in the text.

We have added the concentrations: “300 µM calindol and 100 µM EGCG”.

Reviewer #2 (Recommendations for the authors):

I have included all of my recommendations in the public review section.

Reviewer #3 (Recommendations for the authors):

Although the definitions of 'thickness', 'mouthfulness' and 'continuity' have been revised very helpfully in the Introduction, 'mouthfulness' reappears at other points in the MS e.g., Page 4, Results, Line 3; Page 9, Line 3. It is best replaced by the new definition in these other locations too.

We wish to clarify that our revised text stated, “…to clarify that kokumi attributes are inherently gustatory, in the present study we use the terms ‘intensity of whole complex tastes (rich flavor with complex tastes)’ instead of ‘thickness,’ ‘mouthfulness (spread of taste and flavor throughout the oral cavity)’ instead of ‘continuity,’ and ‘persistence of taste (lingering flavor)’ instead of ‘continuity.’” The term “mouthfulness” was retained in our text, though we provided a more specific explanation. In the re-revised version, we have added “(spread of taste in the oral cavity)” immediately after “mouthfulness.”

I doubt that many scientific readers will be familliar with the term 'intragemmal nerve fibres' (Page 8, Line 4). It is used appropriately but it would be helpful to briefly define/explain it.

We have added an explanation as follows:

“… intragemmal nerve fibers, which are nerve processes that extend directly into the structure of the taste bud to transmit taste signals from taste cells to the brain.”

I previously pointed out the overlap between the CaSR's amino acid (AA) and gamma-glutamyl-peptide binding site. I was surprised by the authors' response which appeared to miss the point being made. It was based on the impacts of selected mutations in the receptor's Venus FlyTrap domain (Broadhead JBC 2011) on the responses to AAs and glutathione analogs. The significantly more active analog, S-methylglutathione is of additional interest because, like glutathione itself, it is present in mammalian body fluids. My apologies to the authors for not more carefully explaining this point.

Thank you for this comment. Both CaSR and GPRC6A are recognized as broad-spectrum amino acid sensors; however, their agonist profiles differ. Aromatic amino acids preferentially activate CaSR, whereas basic amino acids tend to activate GPRC6A. For instance, among basic amino acids, ornithine is a potent and specific activator of GPRC6A, while γ-Glu-Val-Gly in addition to amino acids is a high-potency activator of CaSR. It remains unclear how effectively ornithine activates CaSR and whether γ-glutamyl peptides also activate GPRC6A. These questions should be addressed in future studies.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The concept of kokumi, which refers to an enhanced and more delicious flavor of food, has recently generated considerable interest in food science. However, kokumi has not been well studied in gustatory physiology, and the underlying neuroscientific mechanisms remain largely unexplored. Our previous research demonstrated that ornithine (L-ornithine), which is abundant in shijimi clams, enhanced taste preferences in mice. The present study aimed to build on these findings and investigate the mechanisms responsible for kokumi in rats. In two-bottle preference tests, the addition of ornithine, at a low concentration that did not increase the favorability of this substance alone, enhanced the preferences of animals for umami, sweet, fatty, salty, and bitter solutions, with the intake of monosodium glutamate showing the most significant increase. Additionally, a mixture of umami and ornithine synergistically induced significant responses in the chorda tympani nerve, which transmits taste information to the brain from the anterior part of the tongue. The observed preference enhancement and increase in taste-nerve response were abolished by antagonists of the G-protein-coupled receptor family C group 6 subtype A (GPRC6A). Furthermore, immunohistochemical analysis indicated that GPRC6A was expressed in a subset of type II taste cells in rat fungiform papillae. These results provide new insights into flavor-enhancement mechanisms, confirming that ornithine is a kokumi substance and GPRC6A is a novel kokumi receptor.

Revised version: The concept of “kokumi,” which refers to an enhanced and more delicious flavor of food, has recently generated considerable interest in food science. However, kokumi has not been well studied in gustatory physiology, and the underlying neuroscientific mechanisms remain largely unexplored. Our previous research demonstrated that ornithine (L-ornithine), which is abundant in shijimi clams, enhanced taste preferences in mice. The present study aimed to build on these findings and investigate the mechanisms responsible for kokumi in rats. In two-bottle preference tests, the addition of ornithine, at a low concentration that did not increase the favorability of this substance alone, enhanced the animals’ preferences for umami, sweet, fatty, salty, and bitter solutions, with the intake of monosodium glutamate showing the most significant increase. Additionally, a mixture of umami and ornithine synergistically induced significant responses in the chorda tympani nerve, which transmits taste information to the brain from the anterior part of the tongue. The observed preference enhancement and increase in taste-nerve response were abolished by antagonists of the G-protein-coupled receptor family C group 6 subtype A (GPRC6A). Furthermore, immunohistochemical analysis indicated that GPRC6A was expressed in a subset of type II taste cells in rat fungiform papillae. These results provide new insights into flavor-enhancement mechanisms, confirming that ornithine is a kokumi substance and GPRC6A is a novel kokumi receptor.

Introduction

Earlier version: Enjoying delicious food is a fundamental pleasure in daily life and contributes toward both physical and mental well-being. The perception of deliciousness requires the harmonious integration of various sensory experiences, including the five basic tastes (sweet, salty, sour, bitter, and umami), the taste of fat, spices, aroma, texture, temperature, sight, and sound. These inputs are transmitted to the brain via sensory nerves, where sensory modalities and qualitative information are analyzed (1,2). Beyond these cognitive aspects, the brain further processes these sensations to evoke emotional responses, which significantly contribute to the perception of tastiness or unpleasantness (3-5). Therefore, sensory information from substances entering the oral and nasal cavities influences both cognition and emotions.

Revised version: Enjoying delicious food is a fundamental pleasure in daily life and contributes toward both physical and mental well-being. The perception of deliciousness requires the harmonious integration of various sensory experiences, including the five basic tastes (sweet, salty, sour, bitter, and umami), the taste of fat, spices, aroma, texture, temperature, sight, and sound. These inputs are transmitted to the brain via sensory nerves, where sensory modalities and qualitative information are analyzed (1,2). Beyond these cognitive aspects, the brain further processes these sensations to evoke emotional responses, contributing to palatability or unpleasantness (3–5). Therefore, sensory information from substances entering the oral and nasal cavities influences both cognition and emotions.

Earlier version: Conversely, despite having minimal direct taste effects and a lack of clear information for the brain to process, some substances and their receptors can modify taste information at the peripheral level, thereby controlling palatability. Specifically, recent research on taste modification has focused on “kokumi” substances and receptors (6-8). Kokumi is a Japanese word that denotes enhanced tastiness of food comprising complex ingredients. According to Ohsu et al. (9), when kokumi substances such as glutathione or γ-Glu-Val-Gly are added to chicken consommé at low concentrations that do not produce a specific taste of their own, they bind to the calcium-sensing receptor (CaSR) to induce kokumi by enhancing the “thickness,” “mouthfulness,” and “continuity” of the dish.

Revised version: Conversely, despite having minimal direct taste effects and a lack of clear information for the brain to process, some substances and their receptors can modify taste information mainly at the peripheral level, thereby controlling palatability. Specifically, recent research on taste modification has focused on “kokumi” substances and receptors (6–8). Kokumi is a Japanese word that denotes enhanced tastiness of food comprising complex ingredients. According to Ohsu et al. (9), when kokumi substances such as glutathione or γ-Glu-Val-Gly are added to chicken consommé at low concentrations that do not produce a specific taste of their own, they bind to the calcium-sensing receptor (CaSR) to induce kokumi by enhancing the “thickness,” “mouthfulness,” and “continuity” of the dish.

These three attributes were originally introduced by Ueda et al. (10), who translated Japanese terms describing the profound characteristics of kokumi, which are deeply rooted in Japanese culinary culture (8,11,12). However, these simply translated terms have caused global misunderstanding and confusion, because they sound like somatosensory rather than gustatory descriptions. Therefore, to clarify that kokumi attributes are inherently gustatory, in the present study we use the terms “intensity of whole complex tastes (rich flavor with complex tastes)” instead of thickness, “mouthfulness (spread of taste and flavor throughout the oral cavity),” and “persistence of taste (lingering flavor)” instead of continuity.

Several studies have been conducted on CaSR agonists as candidate kokumi substances (13–21), although neuroscientific research on the mechanism of kokumi induction remains limited and underdeveloped (22–25). In contrast, our previous study in mice (26) focused on ornithine, a non-protein amino acid, as a potential kokumi substance and suggested that the G-protein-coupled receptor family C group 6 subtype A (GPRC6A) is responsible for its action, given that ornithine has a high affinity for this receptor (27–29). As this is the only published study to propose GPRC6A as a potential kokumi receptor other than CaSR, this hypothesis needs to be confirmed.

Earlier version: Neuroscientific studies have reported glutathione (30) in mice and γ-Glu-Val-Gly in mice (23) and rats (25) as kokumi substances. However, to our best knowledge, only our prior study in mice has examined ornithine as a kokumi candidate (26). Therefore, the current study aimed to investigate the effects of ornithine on sweet, salty, sour, bitter, umami, and fatty taste enhancement in rats and to clarify the involvement of the GPRC6A receptor. Whether rodents are suitable models for kokumi research is unclear; however, if the expression of kokumi flavor is fundamental to the enjoyment of delicious food, there should be basic mechanisms common to humans and rodents, even if a similar human perception of kokumi is difficult to detect in rats. Elucidating these neuroscientific mechanisms may shed light on the elusive nature of kokumi in humans.

Revised version: Neuroscientific studies (23,25,30) in mice and rats provide evidence that glutathione and y-Glu-Val-Gly activate CaSRs, and modify behavioral responses to other tastants in a way that may correspond to kokumi as experienced by humans. However, to our best knowledge, only our prior study in mice has examined ornithine as a kokumi candidate (26). Therefore, the current study aimed to investigate the effects of ornithine on sweet, salty, sour, bitter, umami, and fatty taste enhancement in rats and to clarify the involvement of the GPRC6A receptor. Whether rodents are suitable models for kokumi research is unclear; however, if the expression of kokumi flavor is fundamental to the enjoyment of delicious food, there should be basic mechanisms common to humans and rodents, even if a similar human perception of kokumi is difficult to detect in rats. Elucidating these neuroscientific mechanisms may shed light on the elusive nature of kokumi in humans.

Results

Earlier version: First, we used human sensory tests to confirm that adding ornithine to miso soup enhances the three elements of kokumi, namely the intensity of whole complex tastes, mouthfulness, and persistence of taste (see Fig. S1 with the supplemental materials and methods).

Revised version: First, we used human sensory tests to confirm that adding ornithine to miso soup enhances the three elements of kokumi, namely the intensity of whole complex tastes, mouthfulness (spread of taste in the oral cavity), and persistence of taste (see Fig. S1 with the supplemental materials and methods).

Two-bottle preference tests in rats

Additive effects of ornithine at different concentrations

To investigate whether ornithine has a favorable taste, we conducted a preference test by comparing distilled water (DW) with aqueous solutions containing various ornithine concentrations (Fig. 1A). Two-way analysis of variance (ANOVA; solution × concentration) revealed a significant main effect of solution [F(1, 50) = 33.93, P < 0.0001], no main effect of concentration [F(4, 50) = 0.145, P > 0.05], and a significant solution-concentration interaction [F(4, 50) = 4.070, P < 0.01]. Post hoc Bonferroni tests showed that the intake of 10 and 30 mM ornithine was significantly greater (P < 0.01) than that of plain DW, indicating a preference for higher concentrations of ornithine. Next, we examined the additive effect of ornithine at the same concentrations on the preference for 0.03 M monosodium glutamate (MSG; Fig. 1B). Two-way ANOVA revealed a significant main effect of solution [F(1, 50) = 175.41, P < 0.0001] and a solution-concentration interaction [F(4, 50) = 8.371, P < 0.0001], with no main effect of concentration [F(4, 50) = 0.899, P > 0.05]. Subsequent Bonferroni analyses indicated that the addition of ornithine at concentrations ranging from 1–30 mM resulted in significantly greater intake (P < 0.001) than that for MSG alone. Based on these results, we used 1 mM ornithine to examine its additive effects in subsequent experiments.

Additive effects of ornithine (Orn) at different concentrations.

(A) Water intake with and without five Orn concentrations. (B) Intake of 0.03 M MSG with and without Orn. Each value represents the mean ± SEM; n = 6. **P < 0.01, ***P < 0.001 (Bonferroni correction). MSG, monosodium glutamate; SEM, standard error of the mean

Additive effects of 1 mM ornithine on intake of different tastants

We examined the intake of eight taste solutions at four concentrations, both alone and in combination with 1 mM ornithine (Fig. 2). The eight solutions were sweet (sucrose), salty (NaCl), sour (citric acid), bitter (quinine hydrochloride [QHCl]) umami (MSG, monopotassium glutamate [MPG], or inosine monophosphate [IMP]), and fatty (Intralipos). Bonferroni’s multiple-comparison analysis revealed that MSG and MPG at all four concentrations were significantly preferred when combined with ornithine (P < 0.05, 0.01, or 0.001; Fig. 2B and C). This was also observed for IMP (Fig. 2A), Intralipos (Fig. 2D), and sucrose (Fig. 2E) at two concentrations, as well as NaCl (Fig. 2F) at one concentration. However, the rats showed no difference in preference for citric acid regardless of the addition of 1 mM ornithine. In contrast, QHCl with ornithine was preferred at three of the four concentrations, suggesting that the aversive taste of QHCl was attenuated by ornithine. Notably, the effects of ornithine differed among the umami solutions; ornithine more effectively increased the preference for MSG and MPG than for IMP.

Additive effects of 1 mM ornithine (Orn) on fluid intake for eight different taste solutions.

Fluid intake with and without Orn is shown for IMP (A), MSG (B), MPG (C), Intralipos (D), sucrose (E), NaCl (F), citric acid (G), and QHCl (H). Each value represents the mean ± SEM; n = 8. *P < 0.05, **P < 0.01, ***P < 0.001 (Bonferroni correction). IMP, inosine monophosphate; MSG, monosodium glutamate; MPG, monopotassium phosphate; QHCl, quinine hydrochloride; SEM, standard error of the mean

Additive effects of ornithine in brief-exposure tests with and without antagonists

To determine whether the enhancement of preference induced by ornithine was caused by an intra-oral event rather than by post-oral consequences, a brief-exposure (or short-term, 10-minute) two-bottle preference test was conducted. We confirmed that the intake of DW with 1 mM ornithine did not differ from that of plain DW (P > 0.05; Fig. 3A). However, the intake of 0.03 M MSG was significantly increased (P < 0.01) by the addition of 1 mM ornithine (Fig. 3B). This preference was maintained in the presence of 60 µM calindol, a GPRC6A antagonist (Fig. 3C), but disappeared at a calindol concentration of 300 µM (Fig. 3D). Similarly, the GPRC6A antagonist epigallocatechin gallate (EGCG) had no effect at a concentration of 30 μM (Fig. 3E) but abolished the ornithine-enhanced preference at 100 μM (Fig. 3F).

Brief-exposure (10-min) two-bottle preference test results for solutions with and without 1 mM ornithine (Orn) and effects of GPRC6A antagonists calindol and EGCG.

(A) Intake of water with and without Orn. (B) Intake of 0.03 M MSG with and without Orn. (C) Intake of MSG with and without Orn in 60 µM calindol. (D) Intake of MSG with and without Orn in 300 µM calindol. (E) Intake of MSG with and without Orn in 30 µM EGCG. (F) Intake of MSG with and without Orn in 100 µM EGCG. Each value represents the mean ± SEM; n = 7. *P < 0.05, **P < 0.01 (paired t-test, two-tailed). MSG, monosodium glutamate; EGCG, epigallocatechin gallate; GPRC6A, G-protein-coupled receptor family C group 6 subtype A; SEM, standard error of the mean

Additive effects of ornithine after chorda tympani transection

To determine how taste information from the anterior part of the tongue influences ornithine-induced taste preferences, we compared intake during the long-term two-bottle preference test before and after transection of the chorda tympani (CT). As shown in Fig. 4A and C, the combination of 1 mM ornithine and 0.03 M MSG was significantly more ingested than MSG alone (P < 0.01, paired t-test, two-tailed) before both sham and experimental treatments. The same preference was observed following the sham-control operation (Fig. 4B). However, in the experimental group, CT transection abolished the increased favorability induced by ornithine supplementation (Fig. 4D).

Effects of bilateral chorda tympani transection (xCT) and control sham operations (Sham OP) on intake of 0.03 M MSG with and without 1 mM ornithine (Orn).

Fluid intake before (A, C) and after (B, D) the operations. Each value represents the mean ± SEM; n = 4. **P < 0.01 (paired t-test, two-tailed). MSG, monosodium glutamate; SEM, standard error of the mean

Taste-nerve recording in rats

Earlier version: The aqueous solutions of ornithine used in the behavioral experiments induced small CT responses, which increased in a dose-dependent manner, as shown by sample recordings (Fig. 5A) and a graphical representation of mean response magnitudes (Fig. 5D). One-way ANOVA revealed significant differences among the responses at different ornithine concentrations [F(4, 15) = 14.660, P < 0.001]. Additionally, the response to 30 mM ornithine survived the Bonferroni correction for multiple testing, which indicated that this response was significantly higher (P < 0.01) than those at lower concentrations. The response to 1 mM ornithine, which was used in the behavioral experiments, was negligible compared to the standard response to NH4Cl. However, when 1 mM ornithine was added to 0.03 M MSG, the response increased significantly (P < 0.05) compared with that for plain MSG (Fig. 5E). A similar additive effect of ornithine was observed when the solution was prepared with 0.1 mM amiloride, a sodium-channel blocker (Fig. 5B and E), suggesting that glutamate, rather than sodium ions, was responsible for this increased response. One-way ANOVA revealed significant differences among the relative response values shown in Fig. 5E [F(3, 16) = 9.174, P < 0.001]. Finally, we examined the effect of calindol on the increased response to MSG with ornithine. As shown by the sample recordings (Fig. 5C) and graphical representation (Fig. 5F), the increased response to MSG with the addition of ornithine (P < 0.01) was no longer observed when calindol was added. One-way ANOVA revealed significant differences among these responses visualized in Fig. 5F [F(2, 9) = 4.690, P < 0.05]. EGCG, another GPRC6A antagonist, showed effects similar to those of calindol (data not shown).

Revised version: The aqueous solutions of ornithine used in the behavioral experiments induced small CT responses, which increased in a dose-dependent manner, as shown by sample recordings (Fig. 5A) and a graphical representation of mean response magnitudes (Fig. 5D). One-way ANOVA revealed significant differences among the responses at different ornithine concentrations [F(4, 15) = 14.660, P < 0.001]. Additionally, the response to 30 mM ornithine survived the Bonferroni correction for multiple testing, which indicated that this response was significantly higher (P < 0.01) than those at lower concentrations. The response to 1 mM ornithine, which was used in the behavioral experiments, was negligible compared to the standard response to NH4Cl. However, when 1 mM ornithine was added to 0.03 M MSG, the response increased significantly (P < 0.05) compared with that for plain MSG (Fig. 5E). A similar additive effect of ornithine was observed when the solution was prepared with 0.1 mM amiloride, a sodium-channel blocker (Fig. 5B and E), suggesting that glutamate, rather than sodium ions, was responsible for this increased response. One-way ANOVA revealed significant differences among the relative response values shown in Fig. 5E [F(3, 16) = 9.174, P < 0.001]. Finally, we examined the effect of calindol on the increased response to MSG with ornithine. As shown by the sample recordings (Fig. 5C) and graphical representation (Fig. 5F), the increased response to MSG with the addition of ornithine (P < 0.01) was no longer observed when 300 µM calindol was added. One-way ANOVA revealed significant differences among these responses visualized in Fig. 5F [F(2, 9) = 4.690, P < 0.05]. EGCG (100 µM), another GPRC6A antagonist, showed effects similar to those of calindol (data not shown).

Sample recordings of chorda tympani (CT) responses and quantitative representations of mean response magnitudes.

Earlier version: (A) Nerve responses to five concentrations of ornithine (Orn). (B) Nerve responses to 0.03 M MSG with and without 1 mM Orn in water and in 0.01 mM amiloride (sodium-channel blocker). (C) Nerve responses to MSG, MSG with Orn, and MSG with Orn in 300 µM calindol. (D-F) Graphical representations of the mean magnitudes of CT responses corresponding to A, B, and C, respectively. Each value represents the mean ± SEM normalized to the response to 0.1 M NH4Cl = 1.0; n = 4 or 5. *P < 0.05, **P < 0.01, ***P < 0.001 (D and F, Bonferroni correction; E, two-tailed paired t-test). MSG, monosodium glutamate; SEM, standard error of the mean

Revised version: (A) Nerve responses to five concentrations of ornithine (Orn). (B) Nerve responses to 0.03 M MSG with and without 1 mM Orn in water and in 0.01 mM amiloride (sodium-channel blocker). (C) Nerve responses to MSG, MSG with Orn, and MSG with Orn in 300 µM calindol. (D-F) Graphical representations of the mean magnitudes of CT responses corresponding to A, B, and C, respectively. Each value represents the mean ± SEM normalized to the response to 0.1 M NH4Cl = 1.0; n = 4 or 5. Note the different scaling of the ordinate in (D) and (E). *P < 0.05, **P < 0.01, ***P < 0.001 (D and F, Bonferroni correction; E, two-tailed paired t-test). MSG, monosodium glutamate; SEM, standard error of the mean

Immunohistochemical localization of GPRC6A in rat taste cells

Immunohistochemical analyses were performed to determine whether GPRC6A was expressed in the taste cells of rat fungiform, foliate, and circumvallate taste buds. In the fungiform papillae, a small number of spindle-shaped taste cells exhibited GPRC6A-immunoreactivity (Fig. 6A). In the foliate and circumvallate papillae, GPRC6A-immunopositive taste cells were barely detectable, and GPRC6A-expressing cells were likely to constitute less than 1% of the total taste cell population in the respective taste papillae (Fig. 6B, C). These results demonstrated that GPRC6A was preferentially located in subpopulations of fungiform taste cells in the rat.

Immunohistochemical localization of GPRC6A in rat taste papillae.

(A) A significant number of spindle-shaped taste cells exhibited intense GPRC6A immunoreactivity in the fungiform papillae. (B, C) GPRC6A-immunopositive taste cells were barely detectable in the foliate (B) or circumvallate (C) papillae. Left panels show GPRC6A in red, middle panels show Nomarski images of the left panels, and right panels show merged images of respective left and middle panels. Scale bars, 50 µm. GPRC6A, G-protein-coupled receptor family C group 6 subtype A

Earlier version: Since rat taste cells comprise several cell types, we examined the colocalization of GPRC6A and cell type-specific markers in the fungiform papillae. A small subpopulation of GPRC6A-immunopositive cells was found to be immunoreactive for IP3R3, a marker for the majority of the type II cell population (Fig. 7A). Approximately 11% of GPRC6A-positive cells overlapped with IP3R3 (9 double-positive cells/80 GPRC6A-positive cells), while approximately 8.3% of IP3R3-positive cells expressed GPRC6A (9 double-positive /109 IP3R3-positive cells). In contrast, GPRC6A-positive cells were unlikely to colocalize with a-gustducin, another marker for a subset of type II cells, in single taste cells (0 double-positive cell/93 GPRC6A-positive cells) (Fig. 7B). Regarding type III cell markers, GPRC6A-positive cells were unlikely to colocalize with 5-hydroxytryptamine (serotonin, 5-HT) in single taste cells (0 double-positive cell/75 GPRC6A-positive cells) (Fig. 7C). As synaptosomal-associated protein 25 kDa (SNAP-25) labels not only type III cells but also the dense network of intragemmal nerve fibers (31), it is difficult to detect SNAP-25-expressing taste cells surrounded by intense SNAP-25-immunoreactivity of the nerve fibers. However, as far as we investigated, no SNAP-25-expressing cells were found among GPRC6A-positive cells examined (0 double-positive cell/104 GPRC6A-positive cells) (Fig. 7D). These results indicate that GPRC6A is preferentially localized in a subpopulation of IP3R3-expressing taste cells in rat fungiform papillae, and suggest that GPRC6A exerts taste-modifying activity in at least some type II cells in the papillae.

Revised version: Since rat taste cells comprise several cell types, we examined the colocalization of GPRC6A and cell type-specific markers in the fungiform papillae. A small subpopulation of GPRC6A-immunopositive cells was found to be immunoreactive for IP3R3, a marker for the majority of the type II cell population (Fig. 7A). Approximately 11% of GPRC6A-positive cells overlapped with IP3R3 (9 double-positive cells/80 GPRC6A-positive cells), while approximately 8.3% of IP3R3-positive cells expressed GPRC6A (9 double-positive /109 IP3R3-positive cells). In contrast, GPRC6A-positive cells were unlikely to colocalize with a-gustducin, another marker for a subset of type II cells, in single taste cells (0 double-positive cell/93 GPRC6A-positive cells) (Fig. 7B). Regarding type III cell markers, GPRC6A-positive cells were unlikely to colocalize with 5-hydroxytryptamine (serotonin, 5-HT) in single taste cells (0 double-positive cell/75 GPRC6A-positive cells) (Fig. 7C). As synaptosomal-associated protein 25 kDa (SNAP-25) labels not only type III cells but also the dense network of intragemmal nerve fibers, which are nerve fibers that extend directly into the structure of the taste bud to transmit taste signals from taste cells to the brain (31), it is difficult to detect SNAP-25-expressing taste cells surrounded by intense SNAP-25-immunoreactivity of the nerve fibers. However, as far as we investigated, no SNAP-25-expressing cells were found among GPRC6A-positive cells examined (0 double-positive cell/104 GPRC6A-positive cells) (Fig. 7D).

Immunohistochemical analysis of colocalization of GPRC6A and cell type-specific markers in taste cells of rat fungiform papilla taste buds.

(A) Some but not all GPRC6A-immunopositive cells exhibited immunoreactivity for IP3R3, a marker of most type II cells. White arrow indicates GPRC6A/IP3R3 double-positive taste cell. (B) α-Gustducin, another marker of a subset of type II cells, was unlikely to colocalize with GPRC6A in single taste cells. (C, D) Neither 5-HT (C) nor SNAP25 (D), specific markers of type III taste cells, were likely to colocalize with GPRC6A in single taste cells. Scale bars, 10 µm. GPRC6A, G-protein-coupled receptor family C group 6 subtype A; IP3R3, inositol 1,4,5-trisphosphate receptor type 3; 5-HT, 5-hydroxytryptamine; SNAP25, synaptosomal-associated protein 25 kDa

Added in the revised version: These results indicate that GPRC6A is preferentially localized in a subpopulation of IP3R3-expressing taste cells in rat fungiform papillae, and suggest that GPRC6A exerts taste-modifying activity in at least some type II cells in the papillae.

Discussion

Earlier version: In this study, we used a 24-hour two-bottle preference test in rats to demonstrate that the addition of 1 mM ornithine (which itself did not exhibit any taste preference) to tastants comprising umami substances, Intralipos (representing fatty taste), sucrose (sweetness), NaCl (saltiness), and QHCl (bitterness) increased the intake of each solution, suggesting enhanced palatability, especially to umami stimuli. Similar increases in ingestion were observed in a short-term two-bottle preference test using MSG supplemented with ornithine, indicating that this effect was due to the taste sensation in the oral cavity rather than post-ingestive effects. Furthermore, the disappearance of ornithine’s effect on MSG following the application of GPRC6A antagonists and transection of the CT nerve, which innervates taste buds in the anterior tongue, suggests that GPRC6A, which was expressed in the taste cells in this area, is involved in this effect. Supporting this finding, the response of the CT nerve to MSG increased with ornithine supplementation, and the action of a GPRC6A antagonist abolished this change. In addition, immunohistochemical staining showed that GPRC6A predominantly appeared in the fungiform papillae of the anterior tongue rather than in the foliate or circumvallate papillae of the posterior tongue. Despite the differences described below, these results align with our previous findings in mice (26). Together with those from our human study (Fig. S1), these results confirm that ornithine is a potential kokumi substance and that GPRC6A is a candidate kokumi receptor. These insights provide basic neuroscientific knowledge for interpreting the observation in humans that adding ornithine to miso soup enhances the three elements of kokumi: intensity of whole complex tastes, mouthfulness, and persistence of taste.

Revised version: In this study, we used a 24-hour two-bottle preference test in rats to demonstrate that the addition of 1 mM ornithine (which itself did not exhibit any taste preference) to tastants comprising umami substances, Intralipos (representing fatty taste), sucrose (sweetness), NaCl (saltiness), and QHCl (bitterness) increased the intake of each solution, suggesting enhanced palatability, especially to umami stimuli. Similar increases in ingestion were observed in a short-term two-bottle preference test using MSG supplemented with ornithine, indicating that this effect was due to the taste sensation in the oral cavity rather than post-ingestive effects. Furthermore, the disappearance of ornithine’s effect on MSG following the application of GPRC6A antagonists and transection of the CT nerve, which innervates taste buds in the anterior tongue, suggests that GPRC6A, which was expressed in the taste cells in this area, is involved in this effect. Supporting this finding, the response of the CT nerve to MSG increased with ornithine supplementation, and the action of a GPRC6A antagonist abolished this change. In addition, immunohistochemical staining showed that GPRC6A predominantly appeared in the fungiform papillae of the anterior tongue rather than in the foliate or circumvallate papillae of the posterior tongue. Despite the differences described below, these results align with our previous findings in mice (26). Together with those from our human study (Fig. S1), these results confirm that ornithine is a potential kokumi substance and that GPRC6A is a candidate kokumi receptor. These insights provide basic neuroscientific knowledge for interpreting the observation in humans that adding ornithine to miso soup enhances the three elements of kokumi: intensity of whole complex tastes, mouthfulness (spread of taste in the oral cavity), and persistence of taste.

Several similarities and differences are apparent when comparing the results obtained with those of our previous study in mice (26). In both studies, the same concentration of ornithine was used for supplementation (1 mM). The addition of ornithine increased the favorability of basic tastes including umami, sweet, fatty, salty, and bitter tastes, with a more pronounced preference for MSG than IMP among the umami substances. However, while a weak increase in citric acid preference was observed in mice, no such effect was identified in rats. The effects of GPRC6A antagonists and the CT-nerve responses were consistent between the two species. However, the expression of GPRC6A differed considerably; in mice, this protein was expressed throughout the tongue, particularly in the posterior part, whereas in rats, it was predominantly localized to the anterior region of the tongue. Finally, another major difference is that while mice avoided high concentrations of ornithine, rats seemed to prefer them.

Earlier version: Regarding the increased preference for the umami substances IMP, MSG, and MPG when supplemented with ornithine, one may naturally interpret this as an ornithine-enhanced umami response. However, the possibility that umami substances enhance the response to ornithine should be considered. Depending on its concentration, ornithine may stimulate at least three receptors, namely GPRC6A, CaSR (32,33), and the conventional amino acid receptor, the T1R1/T1R3 heterodimer (34). At the low concentration of 1 mM used in the present study, GPRC6A may be exclusively stimulated, whereas higher concentrations may also stimulate T1R1/T1R3, thereby transmitting ornithine taste information via this receptor. Compared with GPRC6A, CaSR may have limited to no involvement, since EGCG, a GPRC6A-specific antagonist, significantly abolished the increased preference induced by ornithine. Umami substances may promote the binding of ornithine to these receptors. For example, Tanase et al. (35) showed that the taste intensity of ornithine increased in the presence of IMP in humans. Similarly, Ueda et al. (36) and Dunkel et al. (37) reported that the taste thresholds for glutathione and glutamyl peptides, agonists of CaSR, were reduced for umami solutions. In rats, ornithine has a favorable taste, and its palatability increased with its concentration (Fig. 1A).

Revised version: Regarding the increased preference for the umami substances IMP, MSG, and MPG when supplemented with ornithine, one may naturally interpret this as an ornithine-enhanced umami response. However, the possibility that umami substances enhance the response to ornithine should be considered. Depending on its concentration, ornithine may stimulate at least three receptors, namely GPRC6A, CaSR (32,33), and the conventional amino acid receptor, the T1R1/T1R3 heterodimer (34). At the low concentration of 1 mM used in the present study, GPRC6A may be exclusively stimulated, whereas higher concentrations may also stimulate T1R1/T1R3, thereby transmitting ornithine taste information via this receptor. Compared with GPRC6A, CaSR may have limited to no involvement, since EGCG, a GPRC6A-specific antagonist, significantly abolished the increased preference induced by ornithine. Umami substances may promote the binding of ornithine to these receptors. For example, Tanase et al. (35) showed that the taste intensity of ornithine increased in the presence of IMP in humans. Similarly, Ueda et al. (36) and Dunkel et al. (37) reported that the taste thresholds for glutathione and glutamyl peptides, agonists of CaSR, were reduced by umami solutions. In rats, ornithine has a favorable taste, and its palatability increased with its concentration (Fig. 1A). When added to ornithine, MSG may enhance the response to ornithine, thereby increasing its palatability. This concept also explains the mixed effects of ornithine and MSG observed in mice. The taste of ornithine itself was increasingly disliked as its concentration was raised (26). Therefore, although the addition of ornithine at concentrations below 3 mM increased palatability, its favorability ceased at concentrations of 10 and 30 mM. Taste information sent by ornithine as its concentration increases, whether pleasant or unpleasant, is likely mediated by the conventional amino acid receptor T1R1/T1R3 rather than GPRC6A. Thus, umami substances may facilitate ornithine stimulation of both kokumi and amino acid receptors in a concentration-dependent manner. Ornithine and umami substances interact to produce synergistic effects on palatability (8).

Removed in the revised version: When added to ornithine, MSG may enhance the response to ornithine, thereby increasing its palatability. This concept also explains the mixed effects of ornithine and MSG observed in mice. The taste of ornithine itself was increasingly disliked as its concentration was raised (26). Therefore, although the addition of ornithine at concentrations below 3 mM increased palatability, its favorability ceased at concentrations of 10 and 30 mM. Taste information sent by ornithine as its concentration increases, whether pleasant or unpleasant, is likely mediated by the conventional amino acid receptor T1R1/T1R3 rather than GPRC6A. Thus, umami substances may facilitate ornithine stimulation of both kokumi and amino acid receptors in a concentration-dependent manner. Ornithine and umami substances interact to produce synergistic effects on palatability (8).

Earlier version: In our experiments with rats (present study) and mice (26), adding ornithine to favorable-tasting solutions such as umami, sweet, and fatty solutions increased their tastiness. This can be explained by the enhanced electrophysiological responses to these tastes induced by ornithine, as evidenced by increased CT-nerve responses for MSG in rats and mice and for other tastes in mice. Conversely, the increased palatability of the aversive bitter taste in both species suggests that ornithine suppresses the response to bitterness. Although the addition of ornithine to quinine slightly decreased the CT-nerve response in mice, this difference was not statistically significant. Reports have shown that ornithine (38,39), and arginine (40), which is structurally highly similar to ornithine, both decrease the bitter taste sensation in humans; however, the underlying mechanism is not well understood. We conducted additional experiments in rats using gallate (Fig. S2), which has been reported as an agonist of GPRC6A (40). Similar to ornithine, gallate increased the palatability of MSG and quinine, although this effect was abolished by EGCG. Notably, this suggests that GPRC6A agonists may augment responses to favorable tastes such as umami and sweetness, while inhibiting responses to bitter tastes, leading to an overall enhancement of deliciousness.

Revised version: The essence of kokumi expression is the enrichment of tastes associated with deliciousness (8). In our experiments with rats (present study) and mice (26), adding ornithine to favorable-tasting solutions such as umami, sweet, and fatty solutions increased their tastiness. This can be explained by the enhanced electrophysiological responses to these tastes induced by ornithine, as evidenced by increased CT-nerve responses for MSG in rats and mice and for other tastes in mice. Conversely, the increased palatability of the aversive bitter taste in both species suggests that ornithine suppresses the response to bitterness. Although the addition of ornithine to quinine slightly decreased the CT-nerve response in mice, this difference was not statistically significant.

Added in the revised version: Reports have shown that ornithine (38,39), and arginine (40), which is structurally highly similar to ornithine, both decrease the bitter taste sensation in humans; however, the underlying mechanism is not well understood. We conducted additional experiments in rats using gallate (Fig. S2), which has been reported as an agonist of GPRC6A (40). Similar to ornithine, gallate increased the palatability of MSG and quinine, although this effect was abolished by EGCG. Notably, this suggests that GPRC6A agonists may augment responses to favorable tastes such as umami and sweetness, while inhibiting responses to bitter tastes, leading to an overall enhancement of deliciousness.

The CaSR agonist γ-Glu-Val-Gly similarly enhanced the preference for umami, sweet, and fatty tastes in rats, but showed no effect on the favorability of salty, sour, or bitter tastes (25). Additionally, while ornithine increased the preference for MSG more than that for IMP, γ-Glu-Val-Gly had the opposite effect. Unlike those of ornithine, the effects of γ-Glu-Val-Gly were not altered by severance of the CT nerve, likely because CaSR is expressed in both the anterior and posterior parts of the tongue (42). The results of the present study, along with those of our previous animal experiments, clearly demonstrate that the binding of kokumi substances to their receptors triggers the expression of kokumi.

However, the mechanism by which this binding enhances the response to umami, fatty, sweet, and other tastes remains unknown. Considering our observation that GPRC6A was co-expressed in type II taste cells positive for IP3R3, some of these cells may be affected by GPRC6A activation. However, it is difficult to explain why co-expression with α-gustducin, which binds to taste receptors in type II cells, was not observed. While we did not examine the co-expression of GPRC6A with the T1R or T2R families, Maruyama et al. (23) reported that CaSR is not co-expressed with receptors for umami, sweet, or other tastes in mice. Some form of intercellular communication within taste buds that affects these tastes is possible, although the detailed mechanism requires further research (43). Whether GPRC6A and CaSR colocalize to the same taste cells should also be considered; however, we have not yet identified any taste cells in which these proteins are co-expressed in our ongoing experiments (Fig. S3). In addition to the above, it is noteworthy that kokumi-active γ-glutamyl peptides have a stronger affinity for the T1R1-MSG receptor complex than for the T1R2-sucrose receptor complex, as demonstrated by molecular modeling approaches (44). This suggests that these peptides exhibit a greater umami-enhancing effect without the involvement of kokumi-specific receptors.

Earlier version: While many candidate kokumi substances have been reported (7,13,14,17-21,22), we are currently exploring the underlying neuroscientific mechanisms using CaSR agonists such as glutathione and γ-Glu-Val-Gly, and GPRC6A agonists such as ornithine and gallate. A common finding in our animal experiments is that these substances enhanced umami to a greater extent than other tastes. They also enhanced sweet and fatty tastes, and to a lesser extent, saltiness. The essence of kokumi expression is the enrichment of tastes associated with deliciousness or the drawing out of newly perceived tastes (8). When tastes are individually enhanced and collectively recognized, the “intensity of whole complex tastes (or rich flavor with complex tastes)” associated with kokumi may arise. During this process, additional taste cells in the oral cavity are more strongly stimulated, and this information is sent to the brain. As a result, the brain perceives this input as coming from the entire mouth, leading to the sensation of “mouthfulness.” Regarding the “persistence of taste (lingering flavor),” it is common sensory physiological knowledge that a strong stimulus causes a large response that takes more time to return to baseline, thereby inducing a longer-lasting sensation. Using a kokumi substance to enhance umami, a taste known for its mouthfulness and lingering quality (8,45,46), further intensifies these characteristics. The comparatively prolonged aftertaste of umami is due to the dominance of umami receptors in taste buds located in the grooves of the circumvallate and foliate papillae at the back of the tongue, where umami substances are less likely to be washed away. This is evidenced by stronger and more specific umami responses from the posterior rather than the anterior part of the tongue in mice (47,48) and primates (46,49). Moreover, MSG binds deeply within the Venus flytrap domain of the umami receptor (50), making it difficult to dislodge.

Revised version: While many candidate kokumi substances have been reported (7,13,14,17–21,22), we are currently exploring the underlying neuroscientific mechanisms using CaSR agonists such as glutathione and γ-Glu-Val-Gly, and GPRC6A agonists such as ornithine and gallate. A common finding in our animal experiments is that these substances enhanced umami to a greater extent than other tastes. They also enhanced sweet and fatty tastes, and to a lesser extent, saltiness. CT recordings indicate that the effects of kokumi stimuli on taste enhancement occur primarily in the peripheral taste organs. When taste responses are individually enhanced and the resulting sensory signals are then transmitted to the brain, where they are processed by the central gustatory and flavor systems, the “intensity of whole complex tastes (or rich flavor with complex tastes)” associated with kokumi may arise. Kokumi substances may reduce the threshold concentrations as well as they increase the suprathreshold responses of tastants. Once the threshold concentrations are lowered, additional taste cells in the oral cavity become activated, and this information is transmitted to the brain. As a result, the brain perceives this input as coming from a wider area of the mouth, leading to the sensation of “mouthfulness.” Regarding the “persistence of taste (lingering flavor),” it is common sensory physiological knowledge that a strong stimulus causes a large response that takes more time to return to baseline, thereby inducing a longer-lasting sensation (e.g., 45).

Added in the revised version: Using a kokumi substance to enhance umami, a taste known for its mouthfulness and lingering quality (8,45,46), further intensifies these characteristics. The comparatively prolonged aftertaste of umami is due to the dominance of umami receptors in taste buds located in the grooves of the circumvallate and foliate papillae at the back of the tongue, where umami substances are less likely to be washed away. This is evidenced by stronger and more specific umami responses from the posterior rather than the anterior part of the tongue in mice (47,48) and primates (46,49). Moreover, MSG binds deeply within the Venus flytrap domain of the umami receptor (50), making it difficult to dislodge.

In conclusion, when kokumi substances are present in complex foods containing umami substances, their interaction primarily enriches the umami taste, followed by sweet and fatty tastes, and to a lesser extent salty and possibly other tastes. This enhancement is at the core of richness, which is accompanied by the secondary sensations of mouthfulness and persistence of taste. Since ornithine is abundant in foods such as shijimi clams (Corbicula japonica) (51,52), cheese (53), and shimeji and maitake mushrooms (54,55), we can enjoy the kokumi-enhanced flavor of cuisines that incorporate these ingredients.

Materials and methods

Animals

Eight-week-old male Wistar rats were obtained from Japan SLC (Shizuoka, Japan). The rats were individually housed in home cages within a temperature-(25 °C) and humidity-controlled (60%) room. A 12:12 hour light/dark cycle was followed, with lights on at 6:00 am and experiments conducted during the light cycle. The animals had free access to food (CE-2 Rodent Diet; CLEA Japan, Inc., Tokyo, Japan) and tap water, except during the brief tests described below in which food access was partially restricted. All animal care and experimental procedures were performed in accordance with the National Institutes of Health guidelines. Experimental protocols were approved by the Institutional Animal Care and Use Committee of Kio University (protocol no. H28-01).

Behavioral experiment: two-bottle preference tests

Each rat was trained to drink DW from a stainless-steel spout connected to a plastic bottle. After a 1-week training period, short- and long-term two-bottle preference tests were conducted, in which two bottles were simultaneously presented to each cage. Each stainless-steel spout, designed to minimize dripping with an internal ball, had an inner diameter of 6 mm and was separated from the center of each spout by 5 cm. Each bottle contained the same taste stimulus (or DW), although one of the two was mixed with ornithine (L-ornithine; Kanto Chemical, Tokyo, Japan).

In the long-term test, the two bottles were switched after 24 hours of the 48-hour test session to account for any potential positional preference. In the short-term (or brief exposure) test, the two bottles with spouts 1 cm apart were presented to the animals for 10 minutes after an overnight water deprivation. Thereafter, the animals had free access to water until 6:00 pm, after which the water was removed overnight. The next day, the bottles were switched and the 10-minute test session was repeated. The bottles containing the fluids were weighed before and after testing to measure the intake volume. The total intake volume over 48 hours or 20 minutes was divided by 2 to obtain the intake volume per day or 10 minutes for the long- and short-term tests, respectively.

Six basic taste stimuli were used: sucrose, NaCl, citric acid, QHCl, MSG (all Kanto Chemical), and Intralipos (a parenteral-stable soybean oil emulsion; Otsuka Pharmaceutical Factory, Tokyo, Japan). In addition to MSG, MPG and IMP (both Ajinomoto, Tokyo, Japan) were also used as umami substances. The taste stimuli were dissolved in DW at various concentrations. A range of concentrations of individual tastants was presented to the same group of animals, starting with the lowest concentration. Ornithine was typically added at a concentration of 1 mM, with other concentrations used as required. A sodium-channel blocker, amiloride (Sigma-Aldrich, Tokyo, Japan), was used to reduce the sodium response to MSG.

For GPRC6A antagonists, we used NPS-2143 (ChemScene, Monmouth Junction, NJ, USA), calindol (Cayman Chemical, Ann Arbor, MI, USA), and EGCG (Tokyo Chemical Industry, Tokyo, Japan). NPS-2143 and calindol were dissolved in 99.5% ethanol at a concentration of 0.1% and diluted to various concentrations using DW. EGCG was prepared in DW.

A total of 89 rats was divided into two groups of 6, three groups of 7, and seven groups of 8 for the taste-preference tests. The first two groups were used for long-term exposure tests of: 1) DW vs. DW with different ornithine concentrations, and 2) 0.03 M MSG vs. 0.03 M MSG with different ornithine concentrations, respectively. The next three groups were used for respective brief-exposure tests of: 1) DW vs. DW with 1 mM ornithine and 0.03 M MSG vs. 0.03 M MSG with 1 mM ornithine, 2) MSG vs. MSG with ornithine, both containing 0.002 or 0.005% calindol, and 3) MSG vs. MSG with ornithine, both containing 30 or 100 μM EGCG. The final seven groups underwent long-term tests. Six of these were used to assess different concentrations of IMP, MSG, MPG, sucrose, Intralipos, and NaCl with and without 1 mM ornithine, respectively. The last group was used to test both citric acid and QHCl with and without 1 mM ornithine.

When examining the range of ornithine and tastant concentrations, we presented the solutions to the animals starting with the lowest concentration. If two different tastants were tested in one group, as in the first of the brief-exposure tests and the long-term test for citric acid and QHCl, the order of presentation was counterbalanced. Statistical analyses were performed within groups or between taste stimuli with and without ornithine, but not across different groups or taste stimuli.

Transection of the CT nerve

Earlier version: Eight naïve rats were randomly divided into transection and sham-operation groups (four rats each). They underwent the long-term two-bottle preference test as described above, with the bottles containing 0.03 M MSG with and without 1 mM ornithine. Thereafter, the animals were anesthetized via intraperitoneal injection of a combination anesthetic (0.3 mg/kg medetomidine, 4.0 mg/kg midazolam, and 5.0 mg/kg butorphanol). In the transection group, the ear ossicles through which the CT innervates the taste buds on the anterior part of the tongue were removed bilaterally, whereas in the sham-operation group, the operation was stopped immediately before the removal of ear ossicles. Postoperatively, the rats were injected with penicillin G sodium (100 mg/kg) to prevent infection. After 6 days of recovery, the same rats were subjected to the long-term two-bottle preference test using the same taste stimuli. The mean group preferences were then compared. After the experiment, transection was confirmed by microscopic verification of the loss of taste buds on the tongue.

Revised version: Eight naïve rats were randomly divided into transection and sham-operation groups (four rats each). They underwent the long-term two-bottle preference test as described above, with the bottles containing 0.03 M MSG with and without 1 mM ornithine. Thereafter, the animals were anesthetized via intraperitoneal injection of a combination anesthetic (0.3 mg/kg medetomidine, 4.0 mg/kg midazolam, and 5.0 mg/kg butorphanol). In the transection group, the ear ossicles through which the CT innervates the taste buds on the anterior part of the tongue were removed bilaterally, whereas in the sham-operation group, the operation was stopped immediately before the removal of ear ossicles.

Added in the revised version: Postoperatively, the rats were injected with penicillin G sodium (100 mg/kg) to prevent infection. After 6 days of recovery, the same rats were subjected to the long-term two-bottle preference test using the same taste stimuli. The mean group preferences were then compared. After the experiment, transection was confirmed by microscopic verification of the loss of taste buds on the tongue.

Taste-nerve recordings

Five naïve rats were deeply anesthetized as described above. Each rat was tracheotomized and secured in a head-holder. The left CT nerve was exposed using a lateral approach (56), excised as it exited the tympanic bulla, and dissected from the underlying tissue. The nerve was placed onto a platinum-wire recording electrode (0.1 mm diameter), while an indifferent electrode was placed in contact with a nearby exposed tissue. The responses were processed using a bandpass filter with cutoff frequencies ranging from 40 Hz to 3 kHz and visualized using an oscilloscope (VC11; Nihon Kohden, Tokyo, Japan). The responses were fed to a digitally controlled summator (57). The number of discharges was summed over 500-ms epochs using a spike counter (DSE-345; DIA Medical System, Tokyo, Japan) to derive summated responses. The data were stored on a PC, and the total spikes over the entire 30-second stimulus period were counted using the PowerLab system (PowerLab/4SP; ADInstruments, Bella Vista, NSW, Australia) for quantitative analyses.

Each taste stimulus (3 mL) was applied to the anterior dorsal tongue for 30 seconds, followed by rinsing with DW for at least 60 seconds. The response to each stimulus was expressed relative to the magnitude of responses to 0.1 M NH4Cl. The analyses of CT results were based on the average values from at least 3 repeated trials in individual animals.

Immunohistochemistry

Male Wistar rats (8-12 weeks old) were deeply anesthetized with isoflurane and transcardially perfused with saline followed by 2% paraformaldehyde in 0.1 M phosphate buffer. To label serotonin-accumulating type III taste cells (58), rats were injected with 5-hydroxy-l-tryptophan at 80 mg/kg body weight 1 hour before anesthesia. Their tongues were dissected out, soaked in 20% sucrose/0.01 M phosphate-buffered saline (PBS) overnight at 4 °C, embedded in optimal-cutting-temperature compound (Sakura Finetek, Tokyo, Japan), and cut into 20-µm thick sections using a cryostat. The sections were immersed in 0.01 M PBS containing 0.2% Triton X-100 and incubated overnight at 4 °C with rabbit anti-GPRC6A antibody (1:50; orb385435; Biorbyt, Cambridge, UK) in combination with either goat anti-IP3R3 (1:100; NB100-2545; Novus, Centennial, CO, USA), anti-α-gustducin (1:500; LSB4942; LSBio, Seattle, WA, USA), anti-5-HT (1:100; ab66047; Abcam, Cambridge, UK), or anti-SNAP25 antibody (1:100; ab31281; Abcam) diluted in the same Triton X-100-containing solution. This was followed by Alexa Fluor 594™-conjugated anti-rabbit immunoglobulin G (IgG; 1:500; A-21207; Thermo Fisher Scientific, Waltham, MA, USA) and 488™-conjugated anti-goat IgG secondary antibodies (1:500; A-11055; Thermo Fisher Scientific). After washing, the sections were cover-slipped with Fluoromount (Diagnostic BioSystems, Pleasanton, CA, USA) and imaged using a Nikon A1Rs confocal laser scanning microscope (Nikon, Tokyo, Japan). The specificity of the anti-GPRC6A antibody was verified as previously described (26) (data not shown). IP3R3 is a marker protein in most type II taste cell populations (59). Further, α-gustducin, 5-HT, and SNAP25 are markers of a subset of types II, III, and III taste cells, respectively (26,58,59).

Statistical analyses

Data are presented as the mean ± standard error of the mean (SEM). A Shapiro–Wilk test was performed to confirm that the data were normally distributed. Student’s t-tests (paired, two-tailed) were used to assess statistical differences between two groups. Additionally, to analyze more than three groups, we used a one-way ANOVA or repeated-measures two-way ANOVA with post hoc Bonferroni tests to account for multiple comparisons. All statistical analyses were conducted in SPSS Statistics Version 25.0 (IBM, Armonk, NY, USA). Statistical significance was set at P < 0.05, except when the Bonferroni correction was used.